---
title: 程序的运行
description: 程序的运行
---

## Python初体验：至尊宝的意中人

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills:
  - 2
```

This is an example exercise. **123**

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}
print(123)
```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
